---
title: 'Chapter 2 - Variables, expression and statements'
description: 'In this module we will be covering the basics of programming languages. With a large focus on values, types, variables and operators. We will also cover the computer uses order of operations to make sense of the code.'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 48fd161bd9
xp: 100
```

Prompting User for Input using Python

`@instructions`
Make your code prompt the user for input and have it display like the code below.

Enter your name: Chuck
Hello Chuck

`@hint`
Use the keyword "input" to get input from the user. You can add a string of words to your input to prompt the user.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
name = 
print("Hello " + name)
```

`@solution`
```{python}
name = input("Enter your name: ")
print("Hello " + name)
```

`@sct`
```{python}

```
